// variables1.rs
//
// Make me compile!
//
// Execute `rustlings hint variables1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

//Before correction
//fn main() {
//    x = 5;
//    println!("x has the value {}", x);
//}

//After correction
fn main() {
    let x: i32 = 5;
    println!("x has the value {}", x);
}


//explication 

// la valeur x n'avait pas ete bien declare, ni type.
// dans rust, pour declarer un variable, on utilise "let".