// variables3.rs
//
// Execute `rustlings hint variables3` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

// Before correction

//fn main() {
//    let x: i32;
//    println!("Number {}", x);
//}

// after correction

fn main() {
    let x: i32 = 22;
    println!("Number {}", x);
}

// Explication
// x n'etait pas initialise donc n'avait aucune valeur a afficher