// functions1.rs
//
// Execute `rustlings hint functions1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

// Before correction

//fn main() {
//    call_me();
//}

// After correction

fn call_me() {
    println!("hey bro");
}

fn main() {
    call_me();
}

// Explication
// la fonction "call_me()" qui est applele n'a pas ete declaree